
--- type:NormalExercise lang:r xp:100 skills:1 key:a4ac7bbd00
## <<<New Exercise>>>
dsghsdklgjhsdougsebngsdvhasg

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
---
title       : Chapter 1:  Tim wonders what to type here
description : testing things out

