
--- type:NormalExercise lang:r xp:100 skills:1 key:a4ac7bbd00
## Learning an old dance craze
Exactly my point! Here–six thousand years of recorded history, a history that includes remarkable composers, astonishing symphonies! But what is the one song that half of them sing to their children generation after generation?

*** =instructions

You put your right hand in,

You put your right hand out,

You put your right hand in,

And you shake it all about,

You do the hokey pokey

and you turn yourself around

That what it's all about.

*** =hint
It doesn't mean anything! I have been studying it for seven days! I had the computer analyze it! I swear to you, it does not mean a thing!

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
x <- 3
y <- 6
z <- x * y
```

*** =solution
```{r}
x <- 3
y <- 6
z <- x * y
```

*** =sct
```{r}

```
---
title       : Chapter 1:  Tim wonders what to type here
description : testing things out

